---
title: "Manipulating, analyzing and exporting data with the tidyverse"
author: "Data Carpentry contributors & MSU R Workshop Leadership Team"
output: learnr::tutorial
runtime: shiny_prerendered
css: css/bootstrap.css
---

```{r setup, include=FALSE}
library(learnr)
library(tidyverse)
library(gradethis)
library(lubridate)
knitr::opts_chunk$set(echo = FALSE, exercise.checker = gradethis::grade_learnr,
                      fig.align = "center")

surveys <- read_csv("data/surveys2_subset.csv")
plots <- read_csv("data/plots.csv")
species <- read_csv("data/species.csv")
```


## Learning Objectives

>
> * Describe the purpose of the **`dplyr`** and **`tidyr`** packages.
> * Select certain columns in a data frame with the **`dplyr`** function `select`.
> * Select certain rows in a data frame according to filtering conditions with the **`dplyr`** function `filter` .
> * Link the output of one **`dplyr`** function to the input of another function with the 'pipe' operator `%>%`.
> * Add new columns to a data frame that are functions of existing columns with `mutate`.
> * Use the split-apply-combine concept for producing data summaries.
> * Use `summarize`, `group_by`, and `count` to split a data frame into groups of observations, apply summary statistics for each group, and then combine the results.
> * Describe the concept of a wide and a long table format and for which purpose those formats are useful.
> * Describe what key-value pairs are.
> * Reshape a data frame from long to wide format and back with the `pivot_wider` and `pivot_longer` commands from the **`tidyr`** package.
> * Export a data frame to a .csv file.

------------

## Data Manipulation using **`dplyr`** and **`tidyr`** Introduction

**`dplyr`** is a 
package for making tabular data manipulation easier. It pairs nicely with 
**`tidyr`** which enables you to swiftly convert between different data formats
for plotting and analysis. Note that we're not using "data 
manipulation," but "data wrangling." To us, "data manipulation" is a term that 
captures the event where a researcher manipulates their data (e.g moving 
columns, deleting rows, merging data files) in a __non-reproducible__ manner. 
Whereas, with data wrangling, all of these process are done, but in a 
__reproducible__ manner, such as using an `R` script!

Packages in `R` are basically sets of additional functions that let you do more
stuff. The functions we've been using so far, like `str()` or `data.frame()`,
come built into `R`; packages give you access to more of them. Before you use a
package for the first time you need to install it on your machine, and then you
should import it in every subsequent `R` session when you need it. You should
already have installed the **`tidyverse`** package. This is an
"umbrella-package" that installs several packages useful for data analysis which
work together well such as **`tidyr`**, **`dplyr`**, **`ggplot2`**, 
**`tibble`**, etc.


The **`tidyverse`** package tries to address 3 common issues that arise when
doing data analysis with some of the functions that come with `R`:

1. The results from a base `R` function sometimes depend on the type of data.
2. Using `R` expressions in a non standard way, which can be confusing for new
   learners.
3. Hidden arguments, having default operations that new learners are not aware
   of.

We have seen in our previous lesson that when building or importing a data 
frame, the columns that contain characters (i.e., text) are coerced (=converted)
into the `factor` data type. We had to set **`stringsAsFactors`** to **`FALSE`**
to avoid this hidden argument to convert our data type. 

This time we will use the **`tidyverse`** package to read the data and avoid 
having to set **`stringsAsFactors`** to **`FALSE`**

If we haven't already done so, we can type `install.packages("tidyverse")` 
straight into the console. In fact, it's better to write this in the console 
than in our script for any package, as there's no need to re-install packages 
every time we run the script.

Then, to load the package type:

```{r, message = FALSE}
## load the tidyverse packages -- including dplyr, tidyr, readr, stringr
library(tidyverse)
```


### What are **`dplyr`** and **`tidyr`**?

The package **`dplyr`** provides easy tools for the most common data 
manipulation tasks. It is built to work directly with data frames, with many 
common tasks optimized by being written in a compiled language (C++). An 
additional feature is the ability to work directly with data stored in an 
external database. The benefits of doing this are that the data can be managed 
natively in a relational database, queries can be conducted on that database, 
and only the results of the query are returned.

This addresses a common problem with `R` in that all operations are conducted
in-memory and thus the amount of data you can work with is limited by available
memory. The database connections essentially remove that limitation in that you
can connect to a database of many hundreds of GB, conduct queries on it directly,
and pull back into `R` only what you need for analysis.

The package **`tidyr`** addresses the common problem of wanting to reshape your 
data for plotting and use by different `R` functions. Sometimes we want data sets 
where we have one row per measurement. Sometimes we want a data frame where each
measurement type has its own column, and rows are instead more aggregated groups
- like plots or aquaria. Moving back and forth between these formats is 
nontrivial, and **`tidyr`** gives you tools for this and more sophisticated data
manipulation.

To learn more about **`dplyr`** and **`tidyr`** after the workshop, you may want
to check out this
[handy data transformation with **`dplyr`** cheatsheet](https://github.com/rstudio/cheatsheets/raw/master/data-transformation.pdf) and this [cheatsheet about **`tidyr`**](https://github.com/rstudio/cheatsheets/raw/master/data-import.pdf).

### Presentation of the Survey Data

We are investigating the animal species diversity and weights found within plots at the Data Carpentry study
site. The dataset is stored as a comma separated value (CSV) file.
Each row holds information for a single animal, and the columns represent:

| Column           | Description                        |
|------------------|------------------------------------|
| record\_id       | Unique id for the observation      |
| month            | month of observation               |
| day              | day of observation                 |
| year             | year of observation                |
| plot\_id         | ID of a particular plot            |
| species\_id      | 2-letter code                      |
| sex              | sex of animal ("M", "F")           |
| hindfoot\_length | length of the hindfoot in mm       |
| weight           | weight of the animal in grams      |


We'll read in our data using the `read_csv()` function, from the tidyverse 
package **`readr`**, instead of `read.csv()`.

```{r load-data, exercise=TRUE}
surveys <- read_csv("https://raw.githubusercontent.com/saramannheimer/data-science-r-workshops/master/Data%20Wrangling/AY%202020-2021/data/surveys.csv")
```


You will see the message `Parsed with column specification`, followed by each 
column name and its data type. When you execute `read_csv` on a data file, it 
looks through the first 1000 rows of each column and guesses the data type for 
each column as it reads it into `R`. For example, in this dataset, `read_csv`reads
`weight` as `col_double` (a numeric data type), and `species` as 
`col_character`. You have the option to specify the data type for a column 
manually by using the `col_types` argument in `read_csv`.


```{r inspect-data, exercise=TRUE}
## inspect the data
glimpse(surveys)
```

```{r eval=FALSE, echo=TRUE}
## Preview the data
View(surveys)
```

```{r echo=FALSE}
surveys
```

Notice that the class of the data is now `tbl_df`

This is referred to as a "tibble". Tibbles tweak some of the behaviors of the 
data frame objects we introduced in the previous episode. The data structure is
very similar to a data frame. For our purposes the only differences are that:

1. In addition to displaying the data type of each column under its name, it
   only prints the first few rows of data and only as many columns as fit on one
   screen.
2. Columns of class `character` are never converted into factors.


We're going to learn some of the most common **`dplyr`** functions:

- `select()`: subset columns
- `filter()`: subset rows on conditions
- `mutate()`: create new columns by using information from other columns
- `group_by()` and `summarize()`: create summary statistics on grouped data
- `arrange()`: sort results
- `count()`: count discrete values

## Select, Filter, and Mutate

### Selecting columns and filtering rows

To select columns of a data frame, use `select()`. The first argument
to this function is the data frame (`surveys`), and the subsequent
arguments are the columns to keep.

Modify the following code to select the `plot_id`, `species_id`, and `weight` columns from the `survey` dataset:

```{r select-columns, exercise=TRUE, exercise.eval=TRUE}
select(surveys)
```

```{r select-columns-solution}
select(surveys, plot_id, species_id, weight)
```

To select all columns *except* certain ones, put a "-" in front of
the variable to exclude it.

Modify the following code to select all columns *except* `record_id` and `species_id`:

```{r donot-select-columns, exercise=TRUE, exercise.eval=TRUE}
select(surveys)
```

```{r donot-select-columns-solution}
select(surveys, -record_id, -species_id)
```

This will select all the variables in `surveys` except `record_id`
and `species_id`.

To choose rows based on a specific criteria, use `filter()`:

```{r filter-rows-1995, exercise=TRUE}
filter(surveys, year == 1995)
```

In the code above `==` keeps all rows where the year is 1995. 

Other filtering options include `!=`, which keeps all rows that are **not** a certain criteria, `,` which means **"and"**, and `|` which means **"or"**. Filter can also do `<` for "less than", `>` for "greater than", `<=` for "less than or equal to", and `>=` for "greater than or equal to". We type these last two options the same way we would typically say them.

1.  `!=` example: 

    ```{r filter-exclude-rows-example, exercise=TRUE}
    filter(surveys, year != 1977)
    ```
    
    The code above keeps all rows where the year is not 1995. 
    
    Modify the following code to keep all rows where the `year` is *not* 1995.

    ```{r filter-exclude-rows, exercise=TRUE, exercise.eval=TRUE}
    filter(surveys, year != 1977)
    ```

    ```{r filter-exclude-rows-solution}
    filter(surveys, year != 1995)
    ```

2.  `,` example:
    
    ```{r filter-comma-example, exercise=TRUE}
    filter(surveys, year == 1977 , plot_id == 2)
    ```
    
    The code above keeps all rows where the year is 1977 for plot id 2, i.e. year 1977 and plot 2. The rows meet **both** of these criteria. 
    
    Modify the following code to keep all rows where the `year` is 1995 and the `plot_id` is 3.
  
    ```{r filter-comma, exercise=TRUE, exercise.eval=TRUE}
    filter(surveys, year == 1977 , plot_id == 2)
    ```
    
    ```{r filter-comma-solution}
    filter(surveys, year == 1995 , plot_id == 3)
    ```

3.  `|` example:

    ```{r filter-or-example, exercise=TRUE}
    filter(surveys, year == 1977 | plot_id == 2)
    ```
    
    The code above keeps all rows where the year is 1977 **or** is plot id 2, i.e. year 1977 or plot 2. The rows meet **either** of these criteria but not both. 
    
    Modify the following code to keep all rows where the `year` is 1995 or plot id 3, but not both.
    
    ```{r filter-or, exercise=TRUE}
    filter(surveys, year == 1977, plot_id == 2)
    ```
    
    ```{r filter-or-solution}
    filter(surveys, year == 1995| plot_id == 3)
    ```

4.  `<` example:

    ```{r less-than-example, exercise=TRUE}
    filter(surveys, weight < 8)
    ```
    
    The code above keeps all rows where weight is less than 8. 
    
    Modify the following code to keep all rows where weight is less than 5.
    
    ```{r less-than, exercise=TRUE, exercise.eval=TRUE}
    filter(surveys, weight < 8)
    ```
    
    ```{r less-than-solution}
    filter(surveys, weight < 5)
    ```


5.  `>` example:

    ```{r greater-than-example, exercise=TRUE}
    filter(surveys, hindfoot_length > 30)
    ```
    
    The code above keeps all rows where hindfoot length is greater than 30.
    
    Modify the following code to add a filter to keep all rows where hindfoot length is greater than 35.
    
    ```{r greater-than, exercise=TRUE, exercise.eval=TRUE}
    filter(surveys)
    ```
    
    ```{r greater-than-solution}
    filter(surveys, hindfoot_length > 35)
    ```

### Pipes

What if you want to select and filter at the same time? There are three
ways to do this: use intermediate steps, nested functions, or pipes.

With intermediate steps, you create a temporary data frame and use
that as input to the next function, like this:



```{r pipes-example-1, exercise=TRUE}
surveys2 <- filter(surveys, weight < 5)
surveys_sml <- select(surveys2, species_id, sex, weight)
```

This is readable, but can clutter up your workspace with lots of objects that 
you have to name individually. With multiple steps, that can be hard to keep 
track of.

You can also nest functions (i.e. one function inside of another), like this:

```{r pipes-example-2, exercise=TRUE}
surveys_sml <- select(filter(surveys, weight < 5), species_id, sex, weight)
```

This is handy, but can be difficult to read if too many functions are nested, as
R evaluates the expression from the inside out (in this case, filtering, then 
selecting).

The last option, *pipes*, are a recent addition to `R`. Pipes let you take
the output of one function and send it directly to the next, which is useful
when you need to do many things to the same dataset.  Pipes in `R` look like
`%>%` and are made available via the **`magrittr`** package, installed 
automatically with **`dplyr`**. If you use RStudio, you can type the pipe with
<kbd>Ctrl</kbd> + <kbd>Shift</kbd> + <kbd>M</kbd> if you have a PC or 
<kbd>Cmd</kbd> + <kbd>Shift</kbd> + <kbd>M</kbd> if you have a Mac.

```{r pipes-example-3, exercise=TRUE}
surveys %>%
  filter(weight < 5) %>%
  select(species_id, sex, weight)
```

In the above code, we use the pipe to send the `surveys` dataset first through
`filter()` to keep rows where `weight` is less than 5, then through `select()`
to keep only the `species_id`, `sex`, and `weight` columns. Since `%>%` takes
the object on its left and passes it as the first argument to the function on
its right, we don't need to explicitly include the data frame as an argument
to the `filter()` and `select()` functions any more.

Some may find it helpful to read the pipe like the word "then". For instance,
in the above example, we took the data frame `surveys`, *then* we `filter`ed
for rows with `weight < 5`, *then* we `select`ed columns `species_id`, `sex`,
and `weight`. The **`dplyr`** functions by themselves are somewhat simple,
but by combining them into linear workflows with the pipe, we can accomplish
more complex manipulations of data frames.

If we want to create a new object with this smaller version of the data, we
can assign it a new name:



```{r pipe-rename-example, exercise=TRUE}
surveys_sml <- surveys %>%
  filter(weight < 5) %>%
  select(species_id, sex, weight)

surveys_sml
```



Note that the final data frame is the leftmost part of this expression.

> ### Challenge 1 ![Medium spicy chili pepper.](images/medium.jpg) {.challenge1}
>
>  Using pipes, subset the `surveys` data to include:
>
  -  animals collected on or after 1995 and 
  -  retain only the columns `year`, `sex`, and `weight`.
>

```{r challenge-1, exercise=TRUE, exercise.lines = 7}
## Pipes Challenge:
##  Using pipes, subset the data to include animals collected
##  before 1995, and retain the columns `year`, `sex`, and `weight.`
```

```{r challenge-1-solution}
surveys %>%
     filter(year >= 1995) %>%
     select(year, sex, weight)
```

```{r challenge-1-check}
grade_code()
```

### Mutate

Frequently you'll want to create new columns based on the values in existing
columns, for example to do unit conversions, or to find the ratio of values in 
two columns. For this we'll use `mutate()`.

To create a new column of weight in kg:

```{r create-column, exercise=TRUE}
surveys %>%
  mutate(weight_kg = weight / 1000)
```



You can also create a second new column based on the first new column within the
same call of `mutate()`:



```{r create-two-columns, exercise=TRUE}
surveys %>%
  mutate(weight_kg = weight / 1000,
         weight_lb = weight_kg * 2.2)
```



If this runs off your screen and you just want to see the first few rows, you
can use a pipe to view the `head()` of the data. (Pipes work with non-**`dplyr`**
functions, too, as long as the **`dplyr`** or `magrittr` package is loaded).



```{r create-column-head, exercise=TRUE}
surveys %>%
  mutate(weight_kg = weight / 1000) %>%
  head()
```



The first few rows of the output are full of `NA`s, so if we wanted to remove
those we could insert a `filter()` in the chain:



```{r mutate-filter, exercise=TRUE, exercise.eval=TRUE}
surveys %>%
  filter(!is.na(weight)) %>%
  mutate(weight_kg = weight / 1000) %>%
  head()
```



`is.na()` is a function that determines whether something is an `NA`. The `!`
symbol negates the result, so in the code above we're asking for every row where weight *is not* 
an `NA`.

> ### Challenge 2 ![Spicy chili pepper.](images/spicy.jpg) {.challenge2}
>
>  Create a new data frame from the `surveys` data named `surveys_hindfoot_cm` 
> that meets the following criteria:
>
-  contains only the `species_id` column and 
-  a new column called `hindfoot_cm` containing the `hindfoot_length` values 
   converted to centimeters.
-  In this `hindfoot_cm` column, there are no `NA`s and all values are less than 3.
-  Print out the new data frame.
>
>  **Hint**: think about how the commands should be ordered to produce this data frame!
> 



```{r challenge-2, exercise=TRUE, exercise.lines = 12, exercise.eval=TRUE} 
## Mutate Challenge:
##  Create a new data frame from the `surveys` data named `surveys_hindfoot_cm`
##  that meets the following criteria: 
##  * contains only the `species_id` column and 
##  * a new column called `hindfoot_cm` containing the `hindfoot_length` values 
##    converted to centimeters.
##  * In this `hindfoot_cm` column, there are no `NA`s and all values are less
##    than 3.
##  Print out the new data frame.

##  Hint: think about how the commands should be ordered to produce this data frame!

```

```{r challenge-2-solution}
surveys_hindfoot_cm <- surveys %>%
  filter(!is.na(hindfoot_length)) %>%
  mutate(hindfoot_cm = hindfoot_length/10) %>%
  filter(hindfoot_cm < 3) %>%
  select(species_id, hindfoot_cm)
surveys_hindfoot_cm
```




```{r challenge-2-check}
grade_result(
  pass_if(~ identical(.result, surveys_hindfoot_cm <- surveys %>%
  filter(!is.na(hindfoot_length)) %>%
  mutate(hindfoot_cm = hindfoot_length/10) %>%
  filter(hindfoot_cm < 3) %>%
  select(species_id, hindfoot_cm)))
)
```
 
### Using `lubridate` for dates 

Date-time data can be frustrating to work with in `R`, since `R` commands for 
date-times are generally un-intuitive and change depending on the type of 
date-time object being used. Moreover, the methods we use with date-times must 
be robust to time zones, leap days, daylight savings times, and other time 
related quirks, and `R` lacks these capabilities in some situations. `Lubridate`
makes it easier to do the things `R` does with date-times and possible to do the
things `R` does not.

`Lubridate` has functions that handle easy parsing of times, such as:  

-  `ymd()`
-  `dmy()` 
-  `mdy()` 


```{r lubridate-example1, message = FALSE, exercise=TRUE}
library(lubridate)
today() # Today's date
now() # Today's date, with time and timezone!
```

```{r lubridate-example2, message = FALSE, exercise=TRUE, exercise.lines=20}
surveys_w_days <- surveys %>% 
  mutate(date = ymd(paste(year,
                          month,
                          day,
                          sep = "-")
                    ),
         day_of_week = wday(date, label = TRUE)
         ## Creating a day of the week variable
         ## label = TRUE prints the name, not the level! 
         )

surveys_w_days %>% 
  filter(is.na(date) == TRUE) %>% 
  select(month, day) %>% 
  table()

head(surveys_w_days)
```

> ### Challenge 3 ![Medium spicy chili pepper.](images/medium.jpg) {.challenge3} 
>
-  What dates were unable to be converted? 
-  Why did that happen?  
>

```{r challenge-4-setup}
surveys_w_days <- surveys %>% 
  mutate(date = ymd(paste(year,
                          month,
                          day,
                          sep = "-")
                    ),
         day_of_week = wday(date, label = TRUE)
         ## Creating a day of the week variable
         ## label = TRUE prints the name, not the level! 
         )
```

We can pull off components of dates using a large array of `lubridate` 
functions, such as:  

-  `year()` 
-  `month()` 
-  `mday()` 
-  `hour()` 
-  `minute()` 
-  `second()`

For additional information about `lubridate` visit [the `lubridate` reference website](https://lubridate.tidyverse.org/) or look over the [`lubridate` cheatsheet](https://github.com/rstudio/cheatsheets/raw/master/lubridate.pdf). 

## Character Wrangling   

If we inspect the day of week variable we created in the last code chunk, we'll 
see that it is an ordered (\<ord\>) factor. 

> ### Challenge 4 ![](images/mild.jpg) {.challenge}
> 
>  What are the names of the days of the week taken from the dates?  


```{r challenge-4, exercise=TRUE, exercise.eval=TRUE}

```

```{r  challenge-4-solution}
levels(surveys_w_days$day_of_week)
```


```{r challenge-4-check}
grade_result(
  pass_if(~ identical(.result, levels(surveys_w_days$day_of_week)))
)
```

### The `case_when()` function

We notice that the labels for the days of the week are not necessarily 
what we would like to have for a graphical display of our data. To reword the 
names of the days of the week, we can use the `case_when()` function from 
`dplyr`. 

The `case_when()` function can be thought of as a "generalized form for 
multiple `if_else()` statements." We talked about `ifelse()` statements in 
the _Intermediate R_ workshop, but let's break them down here to review. 

For `case_when()` the inputs are sequences of two-sided formulas. The left hand 
side finds the values that match the case and the right hand side says what 
should be done with these matches. 

Let's look at this in action! 

```{r caseWhen-setup, message=FALSE, exercise.setup="challenge-4-setup"}
surveys_days_full <- surveys_w_days %>% 
  mutate(day_of_week = case_when(day_of_week == "Mon" ~ "Monday", 
                             day_of_week == "Tue" ~ "Tuesday", 
                             day_of_week == "Wed" ~ "Wednesday", 
                             day_of_week == "Thu" ~ "Thursday", 
                             day_of_week == "Fri" ~ "Friday", 
                             day_of_week == "Sat" ~ "Saturday",
                             day_of_week == "Sun" ~ "Sunday")
         )
glimpse(surveys_days_full$day_of_week)
```

```{r caseWhen-example, exercise=TRUE, exercise.lines=10, exercise.setup="challenge-4-setup"}
surveys_days_full <- surveys_w_days %>% 
  mutate(day_of_week = case_when(day_of_week == "Mon" ~ "Monday", 
                             day_of_week == "Tue" ~ "Tuesday", 
                             day_of_week == "Wed" ~ "Wednesday", 
                             day_of_week == "Thu" ~ "Thursday", 
                             day_of_week == "Fri" ~ "Friday", 
                             day_of_week == "Sat" ~ "Saturday",
                             day_of_week == "Sun" ~ "Sunday")
         )
glimpse(surveys_days_full$day_of_week)
```

> ### NOTE:
> 
> If you only want to recode a couple levels of a variable, you can still use
> `case_when()` without specifying the behavior for **ALL** levels. See the
> example below:

```{r caseWhen-example2, exercise=TRUE}
# Create a variable weekday that takes on a value of 0 for Saturday/Sunday
# and 1 otherwise and recodes Friday to missing
surveys_weekday <- surveys_w_days %>%
  mutate(weekday = case_when(day_of_week == "Sat" ~ 0,
                             day_of_week == "Sun" ~ 0,
                             day_of_week == "Fri" ~ as.numeric(NA),
                             TRUE ~ 1))

surveys_weekday %>%
  count(weekday)
```

But, perhaps these days are not in the order that we want them to be in. 

> ### Challenge 5 ![](images/medium.jpg) {.challenge}
> 
> What order did `R` put the days of the week in?
> What data type is `day_of_week` now? 

```{r challenge-5, exercise=TRUE, exercise.setup="caseWhen-setup"}


```

```{r challenge-5-solution}
table(surveys_day_full$day_of_week)
typeof(surveys_day_full$day_of_week)
```

```{r challenge-5-check}
grade_code()
```

There are small differences between character data types and factor data types. 
Typically, `R` uses factors to handle categorical variables, variables that have
a fixed and known set of possible values. Factors are also helpful for 
reordering character vectors to improve display. However, factors are often 
difficult to work with. Enter the `forcats` package, whose goal is to provide a 
suite of tools that solve common problems with factors, including changing the 
order of levels or the values.

The order of the levels `R` chose may not be what we wanted, but we can 
reorder them using the `fct_relevel()` function from the `forcats` package. 
The function takes three arguments: 

1. the data 
2. the factor to be reordered
3. the order of the new levels separated by commas

This process looks like this: 

```{r relevel-setup, message=FALSE, exercise.setup="caseWhen-setup"}
surveys_edited <- surveys_days_full %>% 
  mutate(day_of_week = fct_relevel(day_of_week, 
                                   "Monday", 
                                   "Tuesday", 
                                   "Wednesday", 
                                   "Thursday", 
                                   "Friday", 
                                   "Saturday", 
                                   "Sunday")
  )
glimpse(surveys_edited$day_of_week)
```

```{r relevel-example, exercise=TRUE, exercise.lines=10, exercise.setup="caseWhen-setup"}
surveys_edited <- surveys_days_full %>% 
  mutate(day_of_week = fct_relevel(day_of_week, 
                                   "Monday", 
                                   "Tuesday", 
                                   "Wednesday", 
                                   "Thursday", 
                                   "Friday", 
                                   "Saturday", 
                                   "Sunday")
  )
glimpse(surveys_edited$day_of_week)
```

> ### Challenge 6
> ![](images/medium.jpg)
> Verify that `R` put the days in the order that you specified! 

```{r challenge-6, exercise=TRUE, exercise.eval=TRUE, exercise.setup="relevel-setup"}


```

```{r challenge-6-solution}
levels(surveys_edited$day_of_week)
```

```{r challenge-6-check}
grade_result(
  pass_if(~ identical(.result, levels(surveys_edited$day_of_week)))
)
```

## Split-apply-combine data analysis

Many data analysis tasks can be approached using the *split-apply-combine*
paradigm: split the data into groups, apply some analysis to each group, and
then combine the results. **`dplyr`** makes this very easy through the use of the
`group_by()` function.


### The `summarize()` function

`group_by()` is often used together with `summarize()`, which collapses each
group into a single-row summary of that group.  `group_by()` takes as arguments
the column names that contain the **categorical** variables for which you want
to calculate the summary statistics. So to compute the mean `weight` by sex:



```{r groupby-example, exercise=TRUE, exercise.setup="relevel-setup"}
surveys_edited %>%
  group_by(sex) %>%
  summarize(mean_weight = mean(weight, na.rm = TRUE))
```



You may also have noticed that the output from these calls doesn't run off the
screen anymore. It's one of the advantages of `tbl_df` over data frame.

You can also group by multiple columns:



```{r groupby-example2, exercise=TRUE, exercise.setup="relevel-setup"}
surveys_edited %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight, na.rm = TRUE))
```



When grouping both by `sex` and `species_id`, the last few rows are for animals
that escaped before their sex and body weights could be determined. You may notice
that the last column does not contain `NA` but `NaN` (which refers to "Not a
Number"). To avoid this, we can remove the missing values for weight before we
attempt to calculate the summary statistics on weight. Because the missing
values are removed first, we can omit `na.rm = TRUE` when computing the mean:



```{r groupby-example3, exercise=TRUE}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight))
```



Here, again, the output from these calls doesn't run off the screen
anymore. If you want to display more data, you can use the `print()` function
at the end of your chain with the argument `n` specifying the number of rows to
display:



```{r groupby-example4, exercise=TRUE}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight)) %>%
  print(n = 15)
```



Once the data are grouped, you can also summarize multiple variables at the same
time (and not necessarily on the same variable). For instance, we could add a
column indicating the minimum weight for each species for each sex:



```{r summarize-example1, exercise=TRUE}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight),
            min_weight = min(weight))
```



It is sometimes useful to rearrange the result of a query to inspect the values.
For instance, we can sort on `min_weight` to put the lighter species first:



```{r summarize-example2, exercise=TRUE}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight),
            min_weight = min(weight)) %>%
  arrange(min_weight)
```



To sort in descending order, we need to add the `desc()` function. If we want to
sort the results by decreasing order of mean weight:



```{r summarize-example3, exercise=TRUE}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(sex, species_id) %>%
  summarize(mean_weight = mean(weight),
            min_weight = min(weight)) %>%
  arrange(desc(mean_weight))
```

> ### Challenge 7 Part 1: ![](images/mild.jpg)
>
> Use `group_by()` and `summarize()` to find the mean, min, and max hindfoot
> length for each species (using `species_id`). Also add the number of observations.
> (HINT: see `?n`.)

```{r challenge-7-1, exercise=TRUE, exercise.lines=7, exercise.setup="relevel-example"}

```

```{r challenge-7-1-solution}
surveys_edited %>%
  filter(!is.na(hindfoot_length)) %>%
  group_by(species_id) %>%
  summarize(mean_hf_length = mean(hindfoot_length),
            min_hf_length = min(hindfoot_length),
            max_hf_length = max(hindfoot_length),
            num_obs = n())
```

```{r challenge-7-1-check}
grade_result(
  pass_if(~ identical(.result, surveys_edited %>%
  filter(!is.na(hindfoot_length)) %>%
  group_by(species_id) %>%
  summarize(mean_hf_length = mean(hindfoot_length),
            min_hf_length = min(hindfoot_length),
            max_hf_length = max(hindfoot_length),
            num_obs = n())))
)
```

> ### Challenge 7 Part 2: ![](images/hot.jpg)
>
> What was the heaviest animal measured in each year?
>
> Return the columns `year`, `genus`, `species_id`, and `weight`.

```{r challenge-7-2, exercise=TRUE, exercise.lines=7, exercise.setup="relevel-example"}

```

```{r challenge-7-2-solution}
surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(year) %>%
  summarize(max_weight = max(weight)) %>%
  select(year, genus, species_id, weight)
```

```{r challenge-7-2-check}
grade_result(
  pass_if(~ identical(.result, surveys_edited %>%
  filter(!is.na(weight)) %>%
  group_by(year) %>%
  summarize(max_weight = max(weight)) %>%
  select(year, genus, species_id, weight)))
)
```

### Counting

When working with data, we often want to know the number of observations found
for each factor or combination of factors. For this task, **`dplyr`** provides
`count()`. For example, if we wanted to count the number of rows of data for
each sex, we would do:



```{r count-example1, exercise=TRUE}
surveys_edited %>%
    count(sex) 
```



The `count()` function is shorthand for something we've already seen: grouping 
by a variable, and summarizing it by counting the number of observations in that
group. In other words, `surveys %>% count()` is equivalent to:  



```{r count-example2, exercise=TRUE}
surveys_edited %>%
    group_by(sex) %>%
    summarize(count = n())
```



For convenience, `count()` provides the `sort` argument:  



```{r count-example3, exercise=TRUE}
surveys_edited %>%
    count(sex, sort = TRUE) 
```



The previous example shows the use of `count()` to count the number of rows/observations 
for *one* factor (i.e., `sex`). 
If we wanted to count *combination of factors*, such as `sex` and `species`, 
we would specify the first and the second factor as the arguments of `count()`:



```{r count-example4, exercise=TRUE}
surveys_edited %>%
  count(sex, species_id) 
```



With the above code, we can proceed with `arrange()` to sort the table 
according to a number of criteria so that we have a better way to compare groups. 
For instance, we might want to arrange the table above in (i) an alphabetical 
order of the levels of the species and (ii) in descending order of the count:



```{r count-example5, exercise=TRUE}
surveys_edited %>%
  count(sex, species_id) %>%
  arrange(species_id, desc(n))
```



From the table above, we may learn that, for instance, there are 75 observations
of the *albigula* species (`species_id` = "NL") that are not specified for its sex (i.e. `NA`), 
the counts for *NL* are on page 3 of the table.



> ### Challenge 8 Part 1: ![](images/mild.jpg)
>
> 1. How many animals were caught in each plot (`plot_id`) surveyed?

```{r challenge-8-p-1, exercise=TRUE, exercise.lines = 12, exercise.eval=TRUE}
## Count Challenges:
##  1. How many animals were caught in each `plot_type` surveyed?

```

```{r challenge-8-p-1-solution}
surveys_edited %>%
        count(plot_type) 
```

```{r challenge-8-p-1-check}
grade_result(
  pass_if(~ identical(.result, surveys_edited %>% count(plot_type) ))
)
```

> ### Challenge 8 Part 2: ![](images/medium.jpg)
>
> 2. Use `group_by()` and `summarize()` to find the mean, min, and max hindfoot
> length for each species (using `species_id`). Also add the number of
> observations (hint: see `?n`).

```{r challenge-8-p-2, exercise=TRUE, exercise.lines = 12, exercise.eval=TRUE}
##  2. Use `group_by()` and `summarize()` to find the mean, min, and max
## hindfoot length for each species (using `species_id`). Also add the number of
## observations (hint: see `?n`).
```

```{r challenge-8-p-2-solution}
surveys_edited %>%
    filter(!is.na(hindfoot_length)) %>%
    group_by(species_id) %>%
    summarize(
            mean_hindfoot_length = mean(hindfoot_length),
            min_hindfoot_length = min(hindfoot_length),
            max_hindfoot_length = max(hindfoot_length),
            n = n()
            )
```

```{r challenge-8-p-2-check}
grade_result(
  pass_if(~ identical(.result, surveys_edited %>%
    filter(!is.na(hindfoot_length)) %>%
    group_by(species_id) %>%
    summarize(
            mean_hindfoot_length = mean(hindfoot_length),
            min_hindfoot_length = min(hindfoot_length),
            max_hindfoot_length = max(hindfoot_length),
            n = n()
            )))
)
```


> ### Challenge 8 Part 3: ![](images/hot.jpg)
> 3. What was the heaviest animal measured in each year? Return the columns 
> `year`, `species_id`, and `weight`.

```{r challenge-8-p-3, exercise=TRUE, exercise.lines = 12, exercise.eval=TRUE}
##  3. What was the heaviest animal measured in each year? Return the
##  columns `year`, `genus`, `species_id`, and `weight`.
```

```{r challenge-8-p-3-solution}
surveys_edited %>%
    filter(!is.na(weight)) %>%
    group_by(year) %>%
    filter(weight == max(weight)) %>%
    select(year, species_id, weight) 
```

```{r challenge-8-p-3-check}
grade_result(
  pass_if(~ identical(.result, surveys_edited %>%
    filter(!is.na(weight)) %>%
    group_by(year) %>%
    filter(weight == max(weight)) %>%
    select(year, species_id, weight) ))
)
```

## Relational Data with `dplyr`   

It is rare that ecological data analyses involve only a single table of data. 
More typically, you have multiple tables of data, describing different aspects 
of your study. When you embark on analyzing your data, these different data 
tables need to be combined. Collectively, multiple tables of data are called 
*relational data*, as the data tables are not independent, rather they relate to
each other. 

Relations are defined between a pair of data tables. There are three families of
joining operations: mutating joins, filtering joins, and set operations. Today 
we will focus on mutating joins. 

The `survey` data have two other data tables they are related to: `plots` and 
`species`. Load in these data and inspect them to get an idea of how they relate
to the `survey` data we've been working with. 

```{r load-data2p1, exercise=TRUE}
plots <- read_csv("https://raw.githubusercontent.com/saramannheimer/data-science-r-workshops/master/Data%20Wrangling/AY%202020-2021/data/plots.csv")

head(plots)
```



Table: Columns in the `plots.csv` file:

| Column           | Description                        |
|------------------|------------------------------------|
| plot\_id         | ID of a particular plot            |
| plot\_type       | type of plot                       |


```{r load-data2p2, exercise=TRUE}
species <- read_csv("https://raw.githubusercontent.com/saramannheimer/data-science-r-workshops/master/Data%20Wrangling/AY%202020-2021/data/species.csv")

head(species)
```

Table: Columns in the `species.csv` file:

| Column           | Description                        |
|------------------|------------------------------------|
| species\_id      | 2-letter code                      |
| genus            | genus of animal                    |
| species          | species of animal                  |
| taxon            | e.g. Rodent, Reptile, Bird, Rabbit |

The variables used to connect a pair of tables are called *keys*. A key is a 
variable that uniquely identifies an observation in that table.  What are the 
keys for each of the three data tables? (hint: What combination of variables 
uniquely identifies a row in that data frame?) 

```{r key-quiz}
quiz(
  question("What is the key for the plots data table?",
    answer("plot_id",correct=TRUE),
    answer("plot_type")
  ),
  question("What is the key for the species data table?",
    answer("species_id", correct = TRUE),
    answer("genus"),
    answer("species"),
    answer("taxa")
  ),
  question("What is the key for the surveys data table?",
    answer("record_id", correct = TRUE),
    answer("month"),
    answer("day"),
    answer("year"),
    answer("plot_id"),
    answer("species_id"),
    answer("sex"),
    answer("hindfoot_length"),
    answer("weight")
  )
)
```

There are two types of keys:  

- A *primary key* uniquely identifies an observation in its own table.  
- A *foreign key* uniquely identifies an observation in another table.  

A primary key and the corresponding foreign key form a *relation* between the 
two data tables. These relations are typically many-to-one, though they can be 
1-to-1. For example, there are many rodents captured that are of one species_id,
hence a many-to-one relationship.  

For me, the easiest way to think about the relationships between the different 
data tables is to draw a picture:  

```{r fig.cap="Relations of survey data tables"}
knitr::include_graphics("images/relations.jpg")
```

 

### Joining Relational Data  

The tool that we will be using is called a *mutating join*. A mutating join is 
how we can combine variables from two tables. The join matches observations by 
their keys, and then copies variables from one table to the other. Similar to 
`mutate()` these join functions add variables to the right of the existing data 
frame, hence their name. There are two types of mutating joins, the inner join 
and the outer join.    

### Inner Join 

The simplest join is an *inner join*, which creates a pair of observations 
whenever their keys are equal. This join will output a new data frame that 
contains the key, the values of `x`, and the values of `y`. Importantly, this 
join deletes observations that do not have a match. 

```{r fig.cap="Wickham, H. & Grolemund, G. (2017) *R for Data Science*. Sebastopol, California: O'Reilly."}
knitr::include_graphics("images/inner.jpg")
```


### Outer Join  

While an inner join only keeps observations with keys that appear in both 
tables, an *outer join* keeps observations that appear in *at least one* of the 
data tables. When joining `x` with `y`, there are three types of outer join:  

- A *left join* keeps all of the observations in `x`.   
- A *right join* keeps all of the observations in `y`.  
- A *full join* keeps all of the observations in both `x` and `y`. 

```{r fig.cap="Wickham, H. & Grolemund, G. (2017) *R for Data Science*. Sebastopol, California: O'Reilly."}
knitr::include_graphics("images/joins.jpg")
```

The left join is the most common, as you typically have a data frame (`x`) that 
you wish to add additional information to (the contents of `y`). This join will 
preserve the contents of `x`, even if there is not a match for them in `y`. 

 

### Joining `surveys_edited` Data  

To join the `surveys_edited` data with the `plots` data and `species` data, we will need
two join statements. As we are interested in adding this information to our 
already existing data frame, `surveys_edited`, a left join is the most appropriate.  



```{r joins-example, exercise=TRUE}
combined <- surveys_edited %>% 
              left_join(plots, by = "plot_id") %>%  # adding the type of plot  
              left_join(species, by = "species_id") # adding the genus, species, and taxa 
glimpse(combined)
```



## Reshaping Data

Data Carpentry's [spreadsheet
lesson (link)](https://datacarpentry.org/spreadsheet-ecology-lesson/01-format-data/),
discusses how to structure our data leading to the four rules defining a tidy
dataset:

1. Each variable has its own column
2. Each observation has its own row
3. Each value must have its own cell
4. Each type of observational unit forms a table

![](images/tidy-data.png)

Here we examine the fourth rule: Each type of observational unit forms a table.

In `surveys_edited` , the rows of `surveys_edited` contain the values of variables associated
with each record (the unit), values such as the weight or sex of each animal 
associated with each record. What if instead of comparing records, we 
wanted to compare the different mean weight of each genus between plots? 
(Ignoring `plot_type` for simplicity).

We'd need to create a new table where each row (the unit) is comprised of values
of variables associated with each plot. In practical terms this means the values
in `genus` would become the names of column variables and the cells would 
contain the values of the mean weight observed on each plot.

Having created a new table, it is therefore straightforward to explore the 
relationship between the weight of different genera within, and between, the
plots. The key point here is that we are still following a tidy data structure,
but we have **reshaped** the data according to the observations of interest:
average genus weight per plot instead of recordings per date.

The opposite transformation would be to transform column names into values of
a variable.

We can do both these of transformations with two `tidyr` functions, 
`pivot_longer()` and `pivot_wider()`.

### Pivoting to a Wider Table

`pivot_wider()` takes three principal arguments:

1. the data  
2. the column whose values will become new column names.  
3. the column whose values will fill the new columns.  

Further arguments include `fill` which, if set, fills in missing values with 
the value provided.

![](images/long_to_wide.JPG)

Let's use `pivot_wider()` to transform surveys to find the mean weight of each 
genus in each plot over the entire survey period. We use `filter()`,
`group_by()` and `summarize()` to filter our observations and variables of 
interest, and create a new variable for the `mean_weight`. We use the pipe as 
before too.

```{r gw-setup, message=FALSE}
combined <- surveys_edited %>% 
              left_join(plots, by = "plot_id") %>%  # adding the type of plot  
              left_join(species, by = "species_id") # adding the genus, species, and taxa 
```

```{r gw-example, exercise=TRUE, exercise.setup="gw-setup"}
surveys_gw <- combined %>%
  filter(!is.na(weight)) %>%
  group_by(plot_id, genus) %>%
  summarize(mean_weight = mean(weight))

glimpse(surveys_gw)
```



This yields `surveys_gw` where the observations for each plot are spread 
across multiple rows, 196 observations of 3 variables.  

Using `pivot_wider()` to pivot on `genus` with values from `mean_weight` this 
becomes 24 observations of 11 variables, one row for each plot. We again use 
pipes:

```{r pivot-wider-setup, message=FALSE, exercise.setup="gw-setup"}
surveys_gw <- combined %>%
  filter(!is.na(weight)) %>%
  group_by(plot_id, genus) %>%
  summarize(mean_weight = mean(weight))
```

```{r pivot-wider-example, exercise=TRUE, exercise.setup="pivot-wider-setup"}
surveys_wide <- surveys_gw %>%
  pivot_wider(names_from = genus, values_from = mean_weight)

glimpse(surveys_wide)
```

> ### Challenge 9: ![](images/spicy.jpg)
> 
> Pivot the `combined` data frame to a wide format, with `year` as columns, 
> `plot_id` as rows, and the number of genera per plot as the values. You will
> need to summarize before reshaping, and use the function `n_distinct()` to 
> get the number of unique genera within a particular chunk of data. It's a 
> powerful function! See `?n_distinct` for more information.
>
> Save the wide dataset as an object, with an intuitive name! Then use
> `glimpse` to take a look at the structure.

```{r challenge-9, exercise=TRUE, exercise.lines=10, exercise.setup="gw-setup"}
## Make a wide data frame by pivoting on year. 
## Fill the values in these columns with the number of genera per plot. 
## Make sure to save the new dataset with an intuitive name!


```

```{r challenge-9-solution}
surveys_wide_genera <- combined %>%
  group_by(plot_id, year) %>%
  summarize(num_genera = n_distinct(genus)) %>%
  pivot_wider(names_from = year, values_from = num_genera)

glimpse(surveys_wide_genera)
```

```{r challenge-9-check}
surveys_wide_genera <- combined %>%
  group_by(plot_id, year) %>%
  summarize(num_genera = n_distinct(genus)) %>%
  pivot_wider(names_from = year, values_from = num_genera)

grade_result(
  pass_if(~ identical(.result, surveys_wide_genera)) 
)
```

### Pivoting to a Longer Table

The opposing situation could occur if we had been provided with data in the
form of `surveys_wide`, where the genus names are column names, but we 
wish to treat them as values of a genus variable instead.

In this situation we are gathering the column names and turning them into a
pair of new variables. One variable represents the column names as values, and
the other variable contains the values previously associated with the column 
names.

`pivot_longer()` takes four principal arguments:

1. the data
2. the columns we wish to pivot into a single column 
3. the name of the new column to create to store the names of each selected column
4. the name of the new column to create to store the data filled in each cell

![](images/wide_to_long.JPG)

To recreate `surveys_gw` from `surveys_wide` we would create a key called
`genus` and value called `mean_weight` and use all columns except `plot_id` for
the key variable. Here we drop `plot_id` column with a minus sign.

```{r pivot-long-setup, message=FALSE, exercise.setup="pivot-wider-setup"}
surveys_wide <- surveys_gw %>%
  pivot_wider(names_from = genus, values_from = mean_weight)
```

```{r pivot-long-example, exercise=TRUE}
surveys_long <- surveys_wide %>%
  pivot_longer(cols = -plot_id, names_to = "genus", values_to = "mean_weight")

glimpse(surveys_long)
```



Note that now the `NA` genera are included in the re-gathered format. Pivoting 
your data to a wide format and then pivoting to a long format can be a useful 
way to balance out a dataset so every replicate has the same composition.

We could also have used a specification for what columns to include. This can be
useful if you have a large number of identifying columns, and it's easier to
specify what to gather than what to leave alone. And if the columns are in a
row, we don't even need to list them all out `-` just use the `:` operator!


```{r pivot-long-example2, exercise=TRUE, exercise.setup="pivot-long-setup"}
surveys_wide %>%
  pivot_longer(cols = Baiomys:Spermophilus, names_to = "genus", 
               values_to = "mean_weight") %>%
  head()
```

> ### Challenge 10 ![](images/spicy.jpg)
>
> Take the `surveys_wide_genera` dataset and use `pivot_longer()` to pivot it to the long 
> format it was in before, so that each row is a unique `plot_id` by `year` 
> combination.
> 
> **HINT**: The year column names look like numbers so you need to use back ticks ("`") to indicate they
> are variables instead of numbers. 


```{r challenge-10-setup}
surveys_wide_genera <- combined %>%
  group_by(plot_id, year) %>%
  summarize(num_genera = n_distinct(genus)) %>%
  pivot_wider(names_from = year, values_from = num_genera)
```

```{r challenge-10, exercise=TRUE, exercise.setup="challenge-10-setup"}
## Now take the surveys_wide_genera dataset, and make it long again, by  
## (re)pivoting on the year columns. 


```

```{r challenge-10-solution}
## Now take the wide dataset, and make it long again, by (re)pivoting on the 
## year columns. 
surveys_wide_genera %>%
  pivot_longer(cols=`1996`:`2002`, names_to = "year",
               values_to = "num_genera")

```

```{r challenge-10-check}
grade_result(
  pass_if(~ identical(.result, surveys_wide_genera %>%
  pivot_longer(cols=`1996`:`2002`, names_to = "year",
               values_to = "num_genera"))) 
)
```


> ### Challenge 11 Part 1: ![](images/hot.jpg)
> 
> The `combined` data set has two measurement columns: `hindfoot_length` and
> `weight`.  This makes it difficult to do things like look at the 
> relationship between mean values of each measurement per year in different
> plot types. 
>
> Let's walk through a common solution for this type of problem.
>
> First, use `pivot_longer()` to create a dataset where we have a names column
> called `measurement` and a values column that takes on the value of either 
> `hindfoot_length` or `weight`. 
>
> **HINT**: You'll need to specify which columns to pivot into longer format!


```{r challenge-11-p-1, exercise=TRUE, exercise.lines=10, exercise.setup="gw-setup"}
## Use pivot_long() to create an even longer dataset.
## Create a column called measurement, containing the hindfoot and weight columns
## And a value column that takes on the value of either of these measurements 
## Hint: You'll need to specify which columns are being used to pivot!



```

```{r challenge-11-p-1-solution}
combined_longer <- combined %>%
  pivot_longer(cols = c(hindfoot_length, weight), names_to = "measurement",
             values_to = "values")

glimpse(combined_longer)
```

```{r challenge-11-p-1-check}
grade_result(
  pass_if(~ identical(.result, combined %>%
  pivot_longer(cols = c(hindfoot_length, weight), names_to = "measurement",
             values_to = "values"))) 
)
```

> ### Challenge 11 Part 2: ![](images/hot.jpg)
> 
> With this new data set, calculate the average of each `measurement` in each
> `year` for each different `plot_type`. 
>
> Then pivot these summaries into a 
> data set with a column for `hindfoot_length` and `weight`. 
>
> **HINT**: This sounds like you want to pivot the data to be a wider format! 

```{r challenge-11-p-2-setup}
combined_longer <- combined %>%
  pivot_longer(cols = c(hindfoot_length, weight), names_to = "measurement",
             values_to = "values")
```

```{r challenge-11-p-2, exercise=TRUE, exercise.lines=20, exercise.setup="challenge-11-p-2-setup"}
## With this new very long data set, calculate the average of each 
## measurement in each year for each different plot_type. 



## Now pivot these summaries into a wide data set. 
## With a columns for hindfoot_length and weight. 
## Filled with the summary values you calculated. 


```

```{r challenge-11-p-2-solution}
measurement_averages <- combined_longer %>%
  group_by(year, plot_type, measurement) %>%
  summarize(avg_measure = mean(values))

measurement_avg_wide <- measurement_averages %>%
  pivot_wider(names_from = measurement, values_from = avg_measure)
```

```{r challenge-11-p-2-check}
grade_result(
  pass_if(~ identical(.result, measurement_averages %>%
  pivot_wider(names_from = measurement, values_from = avg_measure))) 
)
```

## Exporting data

Now that you have learned how to use **`dplyr`** to extract information from
or summarize your raw data, you may want to export these new data sets to share
them with your collaborators or for archival.

Similar to the `read_csv()` function used for reading CSV files into `R`, there is
a `write_csv()` function that generates CSV files from data frames.

Before using `write_csv()`, we are going to create a new folder, `data`,
in our working directory that will store the generated datasets. We don't want
to write generated datasets in the same directory as our raw data. It's good
practice to keep them separate. The `data_raw` folder should only contain the raw,
unaltered data, and should be left alone to make sure we don't delete or modify
it. In contrast, our script will generate the contents of the `data`
directory, so even if the files it contains are deleted, we can always
re-generate them.

```
if(!dir.exists("data")){dir.create("data")}
```

In preparation for our next lesson on plotting, we are going to prepare a
cleaned up version of the dataset that doesn't include any missing data.

Let's start by removing observations of animals for which `weight` and 
`hindfoot_length` are missing, or the `sex` has not been determined:



```{r data-cleanup-exercise, exercise=TRUE}
surveys_complete <- surveys_edited %>%
  filter(!is.na(weight),           # remove missing weight
         !is.na(hindfoot_length),  # remove missing hindfoot_length
         !is.na(sex))                # remove missing sex

glimpse(surveys_complete)
```


Because we are interested in plotting how species abundances have changed
through time, we are also going to remove observations for rare species (i.e.,
that have been observed less than 50 times). We will do this in two steps: first
we are going to create a data set that counts how often each species has been
observed, and filter out the rare species; then, we will extract only the
observations for these more common species:

```{r data-extract-setup, message=FALSE}
surveys_complete <- surveys_edited %>%
  filter(!is.na(weight),           # remove missing weight
         !is.na(hindfoot_length),  # remove missing hindfoot_length
         !is.na(sex))                # remove missing sex
```

```{r data-extract-exercise, exercise=TRUE, exercise.lines=12, exercise.setup="data-extract-setup"}
## Extract the most common species_id
species_counts <- surveys_complete %>%
    count(species_id) %>% 
    filter(n >= 50)

## Only keep the most common species
surveys_complete_subset <- surveys_complete %>%
  filter(species_id %in% species_counts$species_id)
  ## using the relational operator %in% 

glimpse(surveys_complete_subset)
```



To make sure that everyone has the same data set, check that `surveys_complete`
has `r nrow(surveys_complete)` rows and `r ncol(surveys_complete)` columns by
typing `dim(surveys_complete)`.

Now that our data set is ready, we can save it as a CSV file in our `data`
folder.

```{r ready-for-export-setup}
## Extract the most common species_id
species_counts <- surveys_complete %>%
    count(species_id) %>% 
    filter(n >= 50)

## Only keep the most common species
surveys_complete_subset <- surveys_complete %>%
  filter(species_id %in% species_counts$species_id)
  ## using the relational operator %in% 
```

```{r}
dim(surveys_complete_subset)
```

```
write_csv(surveys_complete, path = "data/surveys_complete_subset.csv")
```
